```python
# Initialize Otter
import otter
grader = otter.Notebook("notebook06.ipynb")
```

<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Notebook 06: Charts and Summary Statistics

## Learning Objectives

<div style="background-color:#000; padding:10px;">
</div>

1. Create histograms and boxplots.
2. Computing summary statistics.

## Creating a Histogram for a COVID Data Set  

<div style="background-color:#000; padding:10px;">
</div>

### Exercise 04 📍

The first step is to load in the data, which we will create a dataframe and name it `covid_df`. Then, you will create a histogram for the column `death rate`.  

1. On line 6, assign the `filepath` name to the CSV file `'COVIDbycountry.csv'`. Replace `...` with `'COVIDbycountry.csv'`.
2. Name the dataframe as `covid_df`. Update line 8 (i.e., change `xxx_df` to `covid_df`).    
3. On line 10, update the column name from `???` to `death rate`, and the data frame name from `xxx_df` to `covid_df`.  Notice that when making a histogram, you'd need to specify the column name and data frame.
4. Run the cell to see the histogram for `death rate`.

```python
import pandas as pd
import matplotlib
%matplotlib inline
import seaborn as sns
# Assign filepath to the CSV 'COVIDbycountry.csv'
filepath = ...
# name the dataframe as covid_df.
xxx_df = pd.read_csv(filepath)
# update line 10 to the column name 'death rate' and the dataframe name to 'covid_df'.
deathrate_hg = sns.histplot(x = '???', data = xxx_df)
```

```python
grader.check("exercise_04")
```

## Attribution

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a> and derived from the <a href="#/">APEX STATS Data and Measurement</a> created by David Schuster.

<img src="./by-nc-sa.png" width=100px>

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for exercise_04 (points: None)
```
>>> isinstance(covid_df, pd.DataFrame)
True
```
failure_message: "❌ covid_df is not assigned to a DataFrame."
success_message: "✅ covid_df is assigned to a DataFrame."

